# CS 3110/5110: Data Privacy
## Homework 1

```python
# Load the data and libraries
import pandas as pd
import numpy as np

adult = pd.read_csv('https://github.com/jnear/cs3110-data-privacy/raw/main/homework/adult_with_pii.csv')
adult_pii = adult[['Name', 'DOB', 'SSN', 'Zip', 'Age']]
adult_deid = adult.drop(columns=['Name', 'SSN'])
```

## Question 1 (20 points)

Using the dataframes `adult_pii` and `adult_deid`, write code to conduct a linking attack to recover the names of as many individuals in `adult_deid` as possible. Your solution should be parameterized by the set of columns to use in the attack.

```python
adult
```

Output:
```
Name         DOB          SSN    Zip  Age  \
0       Karrie Trusslove    9/7/1967  732-14-6110  64152   39   
1       Brandise Tripony    6/7/1988  150-19-2766  61523   50   
2          Brenn McNeely    8/6/1991  725-59-9860  95668   38   
3            Dorry Poter    4/6/2009  659-57-4974  25503   53   
4            Dick Honnan   9/16/1951  220-93-3811  75387   28   
...                  ...         ...          ...    ...  ...   
32556       Ardyce Golby  10/29/1961  212-61-8338  41328   27   
32557      Jean O'Connor   6/28/1952  737-32-2919  94735   40   
32558   Reuben Skrzynski    8/9/1966  314-48-0219  49628   58   
32559        Caye Biddle   5/19/1978  647-75-3550   8213   22   
32560  Hortense Hardesty   5/11/2012  690-42-5638  86153   52   

              Workclass  fnlwgt   Education  Education-Num  \
0             State-gov   77516   Bachelors             13   
1      Self-emp-not-inc   83311   Bachelors             13   
2               Private  215646     HS-grad          
```

```python
adult_pii
```

Output:
```
Name         DOB          SSN    Zip  Age
0       Karrie Trusslove    9/7/1967  732-14-6110  64152   39
1       Brandise Tripony    6/7/1988  150-19-2766  61523   50
2          Brenn McNeely    8/6/1991  725-59-9860  95668   38
3            Dorry Poter    4/6/2009  659-57-4974  25503   53
4            Dick Honnan   9/16/1951  220-93-3811  75387   28
...                  ...         ...          ...    ...  ...
32556       Ardyce Golby  10/29/1961  212-61-8338  41328   27
32557      Jean O'Connor   6/28/1952  737-32-2919  94735   40
32558   Reuben Skrzynski    8/9/1966  314-48-0219  49628   58
32559        Caye Biddle   5/19/1978  647-75-3550   8213   22
32560  Hortense Hardesty   5/11/2012  690-42-5638  86153   52

[32561 rows x 5 columns]
```

```python
adult_deid
```

Output:
```
DOB    Zip  Age         Workclass  fnlwgt   Education  \
0        9/7/1967  64152   39         State-gov   77516   Bachelors   
1        6/7/1988  61523   50  Self-emp-not-inc   83311   Bachelors   
2        8/6/1991  95668   38           Private  215646     HS-grad   
3        4/6/2009  25503   53           Private  234721        11th   
4       9/16/1951  75387   28           Private  338409   Bachelors   
...           ...    ...  ...               ...     ...         ...   
32556  10/29/1961  41328   27           Private  257302  Assoc-acdm   
32557   6/28/1952  94735   40           Private  154374     HS-grad   
32558    8/9/1966  49628   58           Private  151910     HS-grad   
32559   5/19/1978   8213   22           Private  201490     HS-grad   
32560   5/11/2012  86153   52      Self-emp-inc  287927     HS-grad   

       Education-Num      Marital Status         Occupation   Relationship  \
0                 13       Never-married       Adm-clerical  Not-in-family   
1    
```

<!-- grade_id: cell-e6389ceace711c32, points: 0, grade: false, solution: true -->
```python
def linking_attack(cols):
    # YOUR CODE HERE
    adult_merge = adult_pii.merge(adult_deid, on=cols)
    
    return adult_merge
adult_merged = linking_attack(['DOB','Zip','Age'])
adult_merged
```

Output:
```
Name         DOB          SSN    Zip  Age  \
0       Karrie Trusslove    9/7/1967  732-14-6110  64152   39   
1       Brandise Tripony    6/7/1988  150-19-2766  61523   50   
2          Brenn McNeely    8/6/1991  725-59-9860  95668   38   
3            Dorry Poter    4/6/2009  659-57-4974  25503   53   
4            Dick Honnan   9/16/1951  220-93-3811  75387   28   
...                  ...         ...          ...    ...  ...   
32556       Ardyce Golby  10/29/1961  212-61-8338  41328   27   
32557      Jean O'Connor   6/28/1952  737-32-2919  94735   40   
32558   Reuben Skrzynski    8/9/1966  314-48-0219  49628   58   
32559        Caye Biddle   5/19/1978  647-75-3550   8213   22   
32560  Hortense Hardesty   5/11/2012  690-42-5638  86153   52   

              Workclass  fnlwgt   Education  Education-Num  \
0             State-gov   77516   Bachelors             13   
1      Self-emp-not-inc   83311   Bachelors             13   
2               Private  215646     HS-grad          
```

<!-- grade_id: cell-251b9f1c6f7b4630, points: 20, grade: true, solution: false -->
```python
# TEST CASES for Question 1

assert len(linking_attack(['Zip'])) == 43191
assert len(linking_attack(['Zip', 'DOB'])) == 32563
assert len(linking_attack(['Zip', 'Age'])) == 32755
```

## Question 2 (5 points)

How many individuals in this dataset are uniquely identified by their Zip code? How many are uniquely identified by their age?

Hint: note that the number of *unique ZIP codes* is **different** from the number of *individuals uniquely identified by ZIP code*.

Hint: you can use the `value_counts` method to count the number of occurences of each value in a series.

<!-- grade_id: cell-0e380ac68dd0c77e, points: 0, grade: false, solution: true -->
```python
def unique_zipcode():
    return len(adult_merged['Zip'].value_counts()[adult_merged['Zip'].value_counts() < 2])

def unique_dob():
    return len(adult_merged['DOB'].value_counts()[adult_merged['DOB'].value_counts() < 2])
```

```python
zipcode_count = unique_zipcode()
zipcode_count
```

Output:
```
23513
```

```python
dob_count = unique_dob()
dob_count
```

Output:
```
7845
```

<!-- grade_id: cell-9b009427c61dbb2b, points: 5, grade: true, solution: false -->
```python
# TEST CASES (INCOMPLETE) for Question 2
# INCOMPLETE means that additional test cases will be used to check your answer during grading

assert unique_zipcode() > 0
assert unique_dob() > 0

```

## Question 3 (10 points)

Write code to determine the `Education-Num` of the individual named Ardyce Golby by performing a differencing attack. Your code should *only* use aggregate data to find Ardyce's education number.

<!-- grade_id: cell-ef28e136429ae8c2, points: 0, grade: false, solution: true -->
```python
def ardyce_education():
    # Differencing attack
    s1 = adult['Education-Num'].sum()
    s2 = adult[adult['Name'] != 'Ardyce Golby']['Education-Num'].sum()
    return s1 - s2
```

```python
ardyce_golby = ardyce_education()
ardyce_golby
```

Output:
```
12
```

<!-- grade_id: cell-d06fb722893a39a5, points: 10, grade: true, solution: false -->
```python
# TEST CASE for Question 3
assert ardyce_education() == 12
```